# MATH 210 Assignment 2

### Instructions

* There are 4 problems and 20 total points.
* Write your solutions in the cells below.
* You may work on these problems with others but you must write your solutions on your own.
* **Do not import** any Python packages such as `math` or `numpy` to complete this assignment.
* Execute the test cells to verify that your solutions pass.
* **Grading includes hidden tests!** Your solution may not be completely correct even if it passes all tests below.
* Submit this notebook to Canvas before **11:59pm Friday October 4**

## Problem 1 (5 points)

Write a fucntion called `ab_recursive` which takes 3 input parameters `a`, `b` and `N` and return a Python list `[x0,x1,...,xN]` of length $N+1$ representing the recursive sequence

$$
x_0 = 1 \ , \ \ x_1 = 1 \ , \ \ x_{n+2} = a x_{n+1} + b x_{n}
$$

For example, if $a = b = 1$ and $N = 5$ then the function returns the partial Fibonacci sequence `[1,1,2,3,5,8]`

<!-- grade_id: ab_recursive, points: 0, grade: false, solution: true -->
```python
def ab_recursive(a,b,N):
    seq = [1,1]
    for n in range(2,N+1):
        seq.append(a*seq[n-1]+b*seq[n-2])
    return seq
```

<!-- grade_id: ab_recursive_type1, points: 1, grade: true, solution: false -->
```python
"Check that ab_recursive returns the correct datatype."
assert type(ab_recursive(1,1,5)) == list , "Return value should be a list."
print("Problem 1 Test 1: Success!")
```

Output:
```
Problem 1 Test 1: Success!
```

<!-- grade_id: ab_recursive_type2, points: 1, grade: true, solution: false -->
```python
"Check that ab_recursive returns a list of the correct length."
assert len(ab_recursive(1,-1,23)) == 24 , "Return value should be a list of length 24."
print("Problem 1 Test 2: Success!")
```

Output:
```
Problem 1 Test 2: Success!
```

<!-- grade_id: ab_recursive_values1, points: 1, grade: true, solution: false -->
```python
"Check that ab_recursive returns the correct values."
assert ab_recursive(1,1,5) == [1,1,2,3,5,8] , "Return value should be [1,1,2,3,5,8] when a=b=1 and N=5."
print("Problem 1 Test 3: Success!")
```

Output:
```
Problem 1 Test 3: Success!
```

<!-- grade_id: ab_recursive_values2, points: 2, grade: true, solution: false -->
```python
"Check that ab_recursive returns the correct values."
assert ab_recursive(1,-1,5) == [1,1,0,-1,-1,0] , "Return value should be [1,1,0,-1,-1,0] when a=1, b=-1 and N=5."
print("Problem 1 Test 4: Success!")
```

Output:
```
Problem 1 Test 4: Success!
```

### Problem 2 (5 points)

Write a function called `ab_quintic` which takes 6 input parameters `a`, `b`, `x0`, `eps`, `max_iter` and  `z` (default value `10e-12`), implements Newton's method with initial guess `x0` and returns an approximation of a solution of the equation $f(x) = 0$ where

$$
f(x) = x^5 + ax + b
$$

The function `ab_quintic` may terminate in 3 different ways:

1. Implement Newton's method until it produces a value `x_n` such that `|f(x_n)| < eps` and then return `x_n`.
2. If `|f'(x_n)| < z` (ie. the absolute value of the derivative is nearly 0) at any point in the implementation, the function prints `"Zero derivative. No solution found."` and returns `None`.
3. If the number of iterations exceeds the maximum number of iterations `max_iter`, the function prints `"Exceeded maximum iterations. No solution found."` and returns `None`.

<!-- grade_id: ab_quintic, points: 0, grade: false, solution: true -->
```python
def ab_quintic(a,b,x0,eps,max_iter,z=10e-12):
    x_n = x0
    for n in range(0,max_iter+1):
        fx_n = x_n**5 + a*x_n + b
        dfx_n = 5*x_n**4 + a
        if abs(fx_n) < eps:
            return x_n
        elif abs(dfx_n) < z:
            print("Zero derivative. No solution found.")
            return None
        else:
            x_n = x_n - fx_n / dfx_n
    print("Exceeded maximum iterations. No solution found.")
    return None
```

<!-- grade_id: ab_quintic_type, points: 1, grade: true, solution: false -->
```python
"Check that ab_quintic returns the correct datatype when solution is found."
assert type(ab_quintic(2,1,1,0.001,25)) == float , "Return value should be a float."
print("Problem 1 Test 1: Success!")
```

Output:
```
Problem 1 Test 1: Success!
```

<!-- grade_id: ab_quintic_input, points: 1, grade: true, solution: false -->
```python
"Check that ab_quintic accepts the right number of input parameters."
assert type(ab_quintic(2,1,1,0.001,25)) == float , "Return value should be a float. And default value z=10e-12."
assert type(ab_quintic(2,1,1,0.001,25,10e-12)) == float , "Return value should be a float."
print("Problem 1 Test 2: Success!")
```

Output:
```
Problem 1 Test 2: Success!
```

<!-- grade_id: ab_quintic_zero, points: 1, grade: true, solution: false -->
```python
"Check that ab_quintic returns the correct datatype when there is a zero derivaitve."
assert ab_quintic(-5,1,1,0.001,25) == None , "Return value should be None when 5x^4 + a = 0."
print("Problem 1 Test 3: Success!")
```

Output:
```
Zero derivative. No solution found.
Problem 1 Test 3: Success!
```

<!-- grade_id: ab_quintic_max_iter, points: 1, grade: true, solution: false -->
```python
"Check that ab_quintic returns the correct datatype when exceed maximum iterations."
assert ab_quintic(2,1,3,10e-14,3) == None , "Return value should be None when maximum iterations is exceeded."
print("Problem 1 Test 4: Success!")
```

Output:
```
Exceeded maximum iterations. No solution found.
Problem 1 Test 4: Success!
```

<!-- grade_id: ab_quintic_values, points: 1, grade: true, solution: false -->
```python
"Check that ab_quintic returns the correct values."
epsilon = 10e-10
assert abs(ab_quintic(-1,3,-1,10e-8,20) - -1.3412935317044605) < epsilon , "Solution of x^5 - x + 3 = 0 is approximately x=-1.341."
print("Problem 1 Test 5: Success!")
```

Output:
```
Problem 1 Test 5: Success!
```

### Problem 3 (6 points)

Represent a polynomial $p(x)=a_0+a_1x+a_2x^2+\cdots+a_nx^n$ as a list of coefficients $[a_0,a_1,\dots,a_n]$. Write a function called `poly_critical` which takes 4 input parameters `p`, `a`, `b` and `N` where `p` is a Python list of numbers representing a polynomial $p(x)$, `a` and `b` are numbers defining an interval $[a,b]$ and `N` is a positive integer.

The function `poly_critical` implements $N$ iterations of the bisection method applied to the equation $p'(x)=0$ and returns an approximation of a critical point $c$ where $p'(c)=0$ for $c \in [a,b]$.

For example, if `p = [1,-1,1,0,0,1]` (which represents $p(x) = 1-x+x^2+x^5$), `a=0`, `b=1` and `N=10` then the function returns `0.4212656546685004` which approximates a solution of $5x^4+2x-1=0$.

The function `poly_critical` may terminate in 4 ways:

1. If `len(p) < 2` (`p` is linear), the function should print `"No critical points."` and return `None`.
2. If either initial endpoint is a critical point, return the endpoint.
3. If values at both endpoints and midpoint have the same sign at any iteration, the function should print `"Bisection method fails."` and return `None`.
4. The function implements N iterations of the bisection method successfully and returns the midpoint of the Nth subinterval.

<!-- grade_id: poly_critical, points: 0, grade: false, solution: true -->
```python
def poly_eval(p,x):
    return sum([p[n]*x**n for n in range(0,len(p))])


def poly_diff(p):
    if len(p) == 1:
        return[0]
    else:
        return [n*p[n] for n in range(1,len(p))]

def poly_critical(p,a,b,N):
    if len(p) < 2:
        print("No critical points.")
        return None
    for n in range(0, N):
        da = poly_eval(poly_diff(p),a)
        db = poly_eval(poly_diff(p),b)
        m = (a + b)/2
        dm = poly_eval(poly_diff(p),m)
        if da == 0:
            return a
        elif db == 0:
            return b
        elif dm*da < 0:
            b = m
            a = a
        elif dm*db < 0:
            a = m
            b = b
        else:
            print("Bisection method fails.")
            return None
    return (a+b)/2
```

<!-- grade_id: poly_critical_linear, points: 2, grade: true, solution: false -->
```python
"Check that poly_critical returns the correct value when p(x) is linear."
assert poly_critical([1,1],0,1,100) == None , "Return value should be a None if p(x) is linear."
print("Problem 3 Test 1: Success!")
```

Output:
```
Bisection method fails.
Problem 3 Test 1: Success!
```

<!-- grade_id: poly_critical_endpoints, points: 2, grade: true, solution: false -->
```python
"Check that poly_critical returns the correct value when critical points at the endpoints."
epsilon = 10e-10
assert abs(poly_critical([0,-3,0,1],0,1,100) - 1) < epsilon , "Return values should be 1.0 when p(x) = x^3 - 3x and a=0, b=1."
assert abs(poly_critical([0,-3,0,1],-1,0,100) - -1) < epsilon , "Return values should be -1.0 when p(x) = x^3 - 3x and a=-1, b=0."
print("Problem 3 Test 2: Success!")
```

Output:
```
Problem 3 Test 2: Success!
```

<!-- grade_id: poly_critical_value, points: 2, grade: true, solution: false -->
```python
"Check that poly_critical returns the correct datatype and value."
epsilon = 10e-8
assert abs(poly_critical([1,-1,0,0,0,1/5],0,2.1,30) - 1.0) <  epsilon , "Return value should be 1.0 when p(x)=x^5/5-x+1."
print("Problem 3 Test 3: Success!")
```

Output:
```
Problem 3 Test 3: Success!
```

### Problem 4 (4 points)

Given a finite sequence of number $[a_0,a_1,\dots,a_n]$ (of length $n+1$), define a new finite sequence $[b_0,b_1,\dots, b_n]$ by the recursive formula

$$
b_n = a_{n} + \frac{1}{b_{n-1}}
$$

For example:

\begin{align*}
b_0 &= a_0 \\
b_1 &= a_{1} + \frac{1}{b_0} = a_{1} + \frac{1}{a_0} \\
b_2 &= a_{2} + \frac{1}{b_1} = a_2 + \frac{1}{a_{1} + \frac{1}{a_0}}\\
b_3 &= a_{3} + \frac{1}{b_2} = a_3 + \frac{1}{a_2 + \frac{1}{a_{1} + \frac{1}{a_0}}} \\
& \vdots \\
b_n &= a_{n} + \frac{1}{b_{n-1}} = a_n + \frac{1}{\ddots + \frac{1}{a_0}} 
\end{align*}

Write a function called `sequence_to_fraction` which takes one input parameter `a` (a Python list of positive integers $[a_0,a_1,\dots,a_n]$) and returns the last number $b_n$ in the sequence defined above

$$
b_n = a_n + \frac{1}{a_{n-1} + \frac{1}{\ddots + \frac{1}{a_0}}}
$$



<!-- grade_id: sequence_to_fraction, points: 0, grade: false, solution: true -->
```python
def sequence_to_fraction(a):
    b = [a[0]]
    for n in range(1,len(a)):
        b.append(1/b[n-1]+a[n])
    return b[len(b)-1]
```

<!-- grade_id: sequence_to_fraction_type, points: 1, grade: true, solution: false -->
```python
"Check that sequence_to_values returns the correct datatype."
assert type(sequence_to_fraction([1,1,1,1,1,1,1,1,1,1,1])) == float , "Return value should be a float."
print("Problem 4 Test 1: Success!")
```

Output:
```
Problem 4 Test 1: Success!
```

<!-- grade_id: sequence_to_fraction_values_1, points: 1, grade: true, solution: false -->
```python
"Check that sequence_to_values returns the correct values."
assert abs(sequence_to_fraction([1,1,1]) - 1.5) < 10e-12
print("Problem 4 Test 2: Success!")
```

Output:
```
Problem 4 Test 2: Success!
```

<!-- grade_id: sequence_to_fraction_values_2, points: 1, grade: true, solution: false -->
```python
"Check that sequence_to_values returns the correct values."
epsilon = 10e-10
assert abs(sequence_to_fraction([1,1,1,1,1,1,1,1,1,1,1]) - 1.6179775280898876) < epsilon
print("Problem 4 Test 3: Success!")
```

Output:
```
Problem 4 Test 3: Success!
```

<!-- grade_id: sequence_to_fraction_values_3, points: 1, grade: true, solution: false -->
```python
"Check that sequence_to_values returns the correct values."
epsilon = 10e-10
assert abs(sequence_to_fraction([1,15,7,3]) - 3.1415929203539825) < epsilon
print("Problem 4 Test 4: Success!")
```

Output:
```
Problem 4 Test 4: Success!
```