# Project 1


### Problem 1 (3 points)

Consider the Taylor series

$$
\frac{1}{2} \ln \left( \frac{1 + x}{1 - x} \right) = \sum_{n=0}^{\infty} \frac{x^{2n+1}}{2n+1} \ , \ x \in (-1,1)
$$

Write a function called `n_sum` which takes 2 input parameters `x` and `N` and returns the sum

$$
\sum_{n=0}^{N} \frac{x^{2n+1}}{2n+1}
$$

<!-- grade_id: n_sum, points: 0, grade: false, solution: true -->
```python
def n_sum(x,N):
    return sum([(x**(2*n+1)/(2*n+1)) for n in range(0,N+1)])
```

<!-- grade_id: n_sum_type, points: 1, grade: true, solution: false -->
```python
"Check that n_sum returns the correct datatype."
assert type(n_sum(1,1)) == float , "Return value should be a float."
print("Problem 1 Test 1: Success!")
```

Output:
```
Problem 1 Test 1: Success!
```

<!-- grade_id: n_sum_values1, points: 1, grade: true, solution: false -->
```python
"Check that n_sum returns the correct values."
epsilon = 10e-14
assert abs(n_sum(1,1) - 4/3) < epsilon , "Value n_sum(1,1) should be 4/3."
print("Problem 1 Test 2: Success!")
```

Output:
```
Problem 1 Test 2: Success!
```

<!-- grade_id: n_sum_values2, points: 2, grade: true, solution: false -->
```python
"Check that n_sum returns the correct values."
epsilon = 10e-14
assert abs(n_sum(0.2,50) - 0.20273255405408222) < epsilon , "Value n_sum(0.2,50) should be 0.20273255405408222"
print("Problem 1 Test 3: Success!")
```

Output:
```
Problem 1 Test 3: Success!
```

### Problem 2 (4 points)

Write a function called `mn_sum` which takes 1 input parameter `N`, and returns the sum

$$
\sum_{n=1}^{N} \sum_{m=1}^{N} \frac{1}{m^2 + n^2}
$$

<!-- grade_id: mn_sum, points: 0, grade: false, solution: true -->
```python
def mn_sum(N):
    return sum([sum([1/(m**2+n**2) for m in range(1,N+1)]) for n in range(1,N+1)])
```

<!-- grade_id: mn_sum_type, points: 1, grade: true, solution: false -->
```python
"Check that mn_sum  returns the correct datatype."
assert type(mn_sum(10)) == float , "Return value should be a float."
print("Problem 2 Test 1: Success!")
```

Output:
```
Problem 2 Test 1: Success!
```

<!-- grade_id: mn_sum_values1, points: 1, grade: true, solution: false -->
```python
"Check that mn_sum returns the correct values."
epsilon = 10e-14
assert abs(mn_sum(1) - 1/2) < epsilon , "Value mn_sum(1) should be 1/2."
print("Problem 2 Test 2: Success!")
```

Output:
```
Problem 2 Test 2: Success!
```

<!-- grade_id: mn_sum_values2, points: 2, grade: true, solution: false -->
```python
"Check that mn_sum returns the correct values."
epsilon = 10e-14
assert abs(mn_sum(2) - (1/2 + 2/5 + 1/8)) < epsilon , "Value mn_sum(2) should be 41/40."
print("Problem 2 Test 3: Success!")
```

Output:
```
Problem 2 Test 3: Success!
```

### Problem 3 (4 points)

Represent a polynomial $p(x)=c_0+c_1x+c_2x^2+\cdots+c_dx^d$ as a list of coefficients $[c_0,c_1,c_2,\dots,c_d]$. Write a function called `poly_int` which takes 3 input parameters `p`, `a` and `b`, where `p` is a Python list of numbers representing a polynomial $p(x)$, and `a` and `b` are numbers. The function returns the value of the definite integral

$$
\int_a^b p(x) \, dx
$$

<!-- grade_id: poly_int, points: 0, grade: false, solution: true -->
```python
def take_int(num,letter,p):
    if num == 0:
        return num
    return p[num-1]*(1/num)*letter**num + take_int(num - 1, letter,p)


def poly_int(p,a,b):
    length = len(p)
    return take_int(length,b,p) - take_int(length,a,p)
    
    
```

<!-- grade_id: poly_int_type, points: 1, grade: true, solution: false -->
```python
# "Check that poly_int returns the correct type."
assert type(poly_int([1,1,1],0,1)) == float , "Return value should be a float."
print("Problem 3 Test 1: Success!")
```

Output:
```
Problem 3 Test 1: Success!
```

<!-- grade_id: poly_int_values1, points: 1, grade: true, solution: false -->
```python
"Check that poly_int returns the correct values."
epsilon = 10e-14
assert abs(poly_int([0,1],0,1) - 1/2) < epsilon , "Value poly_int([0,1],0,1) should be 1/2."
print("Problem 3 Test 2: Success!")
```

Output:
```
Problem 3 Test 2: Success!
```

<!-- grade_id: poly_int_values2, points: 2, grade: true, solution: false -->
```python
"Check that poly_int returns the correct values."
epsilon = 10e-14
assert abs(poly_int([1,1,1,1],0,1) - (1 + 1/2 + 1/3 + 1/4)) < epsilon , "Value poly_int([1,1,1,1],0,1) should be 1 + 1/2 + 1/3 + 1/4."
print("Problem 3 Test 3: Success!")
```

Output:
```
Problem 3 Test 3: Success!
```

### Problem 4  (8 points)

Write a function called `statistics` which takes one input parameter `samples`, a list of numbers, and returns a tuple (of length 7) of statistics:

    (mean, standard deviation, minimum, median, maximum, range, total number of samples)

Recall, if `samples` is the list of numbers $[x_0,x_1,...,x_N]$ then

$$
\begin{align}
\text{mean} & : & \mu = \frac{1}{N+1}\sum_{k=0}^N x_k \\
\text{standard deviation} & : & \sqrt{\frac{\sum_{k=0}^N (\mu - x_k)^2}{N + 1}} \\
\text{range} & : & \max\{x_k\} - \min\{x_k\} \\
\end{align}
$$

The [median](https://en.wikipedia.org/wiki/Median) is the middle value if the number of samples is odd, or the median is the average of the two middle values if the number of samples is even. Use the built-in function `sorted` to sort the list of values. 

Remember:
* Do not import any Python packages
* Do not use Python keywords or built-in function names (such as `max`, `min`, and `range`) as variable names

<!-- grade_id: statistics, points: 0, grade: false, solution: true -->
```python
def calc_range(list):
    return max(list) - min(list)

def calc_mean(list):
    return sum([x/len(list) for x in list])

def calc_stdev(list):
    mean = calc_mean(list)
    return (sum([(mean-x)**2 for x in list])/len(list))**0.5

def calc_median(list):
    stlist = sorted(list)
    if  (len(list) % 2) == 1:
        return stlist[int((len(list)-1)/2)]
    return (stlist[int((len(list)-2)/2)] + stlist[int(len(list)/2)])/2

def statistics(samples):
    return (calc_mean(samples), 
            calc_stdev(samples), 
            min(samples), 
            calc_median(samples), 
            max(samples), 
            calc_range(samples), 
            len(samples))
```

<!-- grade_id: statistics_type, points: 1, grade: true, solution: false -->
```python
"Check that statistics returns the correct datatype."
assert type(statistics([1.3,5.2,3.7])) == tuple , "Return value should be a tuple."
assert len(statistics([1.3,5.2,3.7])) == 7 , "Return value should be a tuple of length 7."
print("Problem 4 Test 1: Success!")
```

Output:
```
Problem 4 Test 1: Success!
```

<!-- grade_id: statistics_mean, points: 1, grade: true, solution: false -->
```python
"Check that statistics returns the correct mean value."
epsilon = 10e-14
assert abs(statistics([1.0,2.0,3.0])[0] - 2.0) < epsilon , "Mean value failed."
print("Problem 4 Test 2: Success!")
```

Output:
```
Problem 4 Test 2: Success!
```

<!-- grade_id: statistics_median, points: 2, grade: true, solution: false -->
```python
"Check that statistics returns the correct median value."
assert abs(statistics([1.0,2.0,3.0])[3] - 2.0) < epsilon , "Median value failed."
print("Problem 4 Test 3: Success!")
```

Output:
```
Problem 4 Test 3: Success!
```

<!-- grade_id: statistics_std, points: 2, grade: true, solution: false -->
```python
"Check that statistics returns the correct standard deviation value."
epsilon = 10e-14
assert abs(statistics([1.0,2.0,3.0])[1] - (2/3)**0.5) < epsilon , "Standard deviation value failed."
print("Problem 4 Test 4: Success!")
```

Output:
```
Problem 4 Test 4: Success!
```

<!-- grade_id: statistics_range, points: 2, grade: true, solution: false -->
```python
"Check that statistics returns the correct range values."
epsilon = 10e-14
assert abs(statistics([1.0,2.0,3.0])[2] - 1.0) < epsilon , "Minimum value failed."
assert abs(statistics([1.0,2.0,3.0])[4] - 3.0) < epsilon , "Maximum value failed."
assert abs(statistics([1.0,2.0,3.0])[5] - 2.0) < epsilon , "Range value failed."
assert statistics([1.0,2.0,3.0])[6] == 3 , "Total value failed."
print("Problem 4 Test 5: Success!")
```

Output:
```
Problem 4 Test 5: Success!
```

```python

```