# Project 4


```python
import numpy as np
import matplotlib.pyplot as plt
import scipy.linalg as la
```

### Problem 1 (5 points)

Write a function called `cval` which takes a 2D NumPy array $A$ representing a square matrix and returns the value

$$
c = \sqrt{\frac{\lambda_{\max}}{\lambda_{\min}}}
$$

where $\lambda_{\max}$ and $\lambda_{\min}$ are the maximum and minimum eigevalues of $A^TA$. Note that the eigenvalues of $A^TA$ are real and non-negative. If $\lambda_{\min} < 10^{-14}$ (ie. $\lambda_{\min} = 0$), then return `np.inf`. If the matrix $A$ is not square, print the statement `"Matric is not square"` and return `None`.

<!-- grade_id: cval, points: 0, grade: false, solution: true -->
```python
def cval(A):
    if np.shape(A)[0] != np.shape(A)[1]:
        print("Matrix is not square")
        return None
    else:
        vals = la.eig(A.T@A)[0].real
        lmax = max(vals)
        lmin = min(vals)
        if lmin < 10**(-14):
            return np.inf
        else:
            return np.sqrt(lmax / lmin)        
```

<!-- grade_id: cval_type, points: 1, grade: true, solution: false -->
```python
"Check cval returns the correct datatype for square matrix."
A = np.random.rand(5,5)
assert type(cval(A)) == np.float64 , "Return value should be a NumPy float."
print("Problem 1 Test 1: Success!")
```

Output:
```
Problem 1 Test 1: Success!
```

<!-- grade_id: cval_none, points: 1, grade: true, solution: false -->
```python
"Check cval returns None for non-square matrix."
A = np.random.rand(3,5)
assert cval(A) == None , "Return value should be None when A is not square."
print("Problem 1 Test 2: Success!")
```

Output:
```
Matrix is not square
Problem 1 Test 2: Success!
```

<!-- grade_id: cval_inf, points: 1, grade: true, solution: false -->
```python
"Check cval returns inf when \lambda_{min} == 0."
A1 = np.ones((4,4))
assert cval(A1) == np.inf , "Return value should be inf when det(A) = 0."
print("Problem 1 Test 3: Success!")
```

Output:
```
Problem 1 Test 3: Success!
```

<!-- grade_id: cval_values, points: 2, grade: true, solution: false -->
```python
"Check cval returns correct values."
epsilon = 10e-10
A1 = np.eye(5)
assert np.abs(cval(A1) - 1.0) < epsilon , "Return value should be 1.0 for the identity matrix."
A2 = np.diag(np.arange(1,9))
assert np.abs(cval(A2) - 8.0) < epsilon , "Return value should be 8.0 for the diagonal matrix with entries [1,2,3,4,5,6,7,8]."
print("Problem 1 Test 4: Success!")
```

Output:
```
Problem 1 Test 4: Success!
```

### Problem 2 (5 points)


Let $n$ be a positive integer and let $h$ be a small positive number. Let $D_n$ be the square matrix of size $n$ such that

\begin{align*}
D_{n; 0,0} &= -\frac{1}{h} \\
D_{n; 0,1} &= \frac{1}{h} \\
D_{n; i,i+1} &= \frac{1}{2h} \text{ for } i=1,\dots,n-2 \\
D_{n; i+1,i} &= -\frac{1}{2h} \text{ for } i=0,\dots,n-3 \\
D_{n; n-1,n-2} &= -\frac{1}{h} \\
D_{n; n-1,n-1} &= \frac{1}{h} \\
\end{align*}

The subscripts $D_{n; i,j}$ indicate the entry at row index $i$ and column index $j$ of the matrix $D_n$. For example, when $n=4$ we have

$$
D_4 = \frac{1}{2h} \left[ \begin{array}{rrrr}
-2 & 2 & 0 & 0 \\
-1 & 0 & 1 & 0 \\
0 & -1 & 0 & 1 \\
0 & 0 & -2 & 2 \\
\end{array} \right]
$$

Write a function called `solve_D` which takes input parameters $n$, $h$ and $b$ (a 2D NumPy array of size $n$ by 1) and returns a solution of the system of equations

$$
\left( D_n^2 - I_n \right) x = b
$$

where $I_n$ is the identity matrix of size $n$. The return value $x$ is a 2D NumPy array of size $n$ by 1. The NumPy function `np.diag` is helpful.

<!-- grade_id: solve_D, points: 0, grade: false, solution: true -->
```python
def solve_D(n,h,b):
    I_n = np.eye(n,n)
    D = np.zeros((n,n))
    D[0,0] = -2
    D[0,1] = 2
    D[n-1,n-2] = -2
    D[n-1,n-1] = 2
    for n in range(0,n-2):
        D[n+1,n+2] = 1
        D[n+1,n] = -1
    D_n = (1 / (2*h))*D
    Dif = D_n@D_n - I_n
    return la.solve(Dif,b)
```

<!-- grade_id: solve_D_type, points: 1, grade: true, solution: false -->
```python
"Check that solve_D returns the correct type."
assert type(solve_D(3,1,np.array([0,1,2]).reshape(3,1))) == np.ndarray, "Return type should be a NumPy array."
print("Problem 2 Test 1: Success!")
```

Output:
```
Problem 2 Test 1: Success!
```

<!-- grade_id: solve_D_shape, points: 1, grade: true, solution: false -->
```python
"Check that solve_D returns array of the correct shape."
n = np.random.randint(2,10)
h = np.random.rand()
b = np.arange(0,n).reshape(n,1)
assert solve_D(n,h,b).shape == (n,1), "Return type should be a 2D NumPy array of size (n,1)."
print("Problem 2 Test 2: Success!")
```

Output:
```
Problem 2 Test 2: Success!
```

<!-- grade_id: solve_D_values, points: 3, grade: true, solution: false -->
```python
"Check that solve_D returns the correct values."
epsilon = 10e-12
assert np.max(np.abs(solve_D(3,1,np.array([1,-2,4]).reshape(3,1)) - np.array([[-5.5],[-2.5],[-8.5]]))) < epsilon
print("Problem 2 Test 3: Success!")
```

Output:
```
Problem 2 Test 3: Success!
```

##  Problem 3 (4 points)

Let $H_n$ be the square matrix of size $n$ such that

$$
H_{n; i,j} = \frac{1}{i + j + 2} \ \  0 \leq i,j < n
$$


The subscripts $H_{n; i,j}$ indicate the entry at row index $i$ and column index $j$ of the matrix $H_n$. For example, when $n=4$ we have

$$
H_4 = \left[ \begin{array}{rrrr}
1/2 & 1/3 & 1/4 & 1/5 \\
1/3 & 1/4 & 1/5 & 1/6 \\
1/4 & 1/5 & 1/6 & 1/7 \\
1/5 & 1/6 & 1/7 & 1/8
\end{array} \right]
$$

Write a function called `solve_H` which takes input parameter $n$ and $b$ (a 2D NumPy array of size $n$ by 1), and returns a solution of the system of equations

$$
\left( H_n^2 - I_n \right) x = b
$$

where $I_n$ is the identity matrix of size $n$. The return value $x$ is a 2D NumPy array of size $n$ by 1.

<!-- grade_id: solve_H, points: 0, grade: false, solution: true -->
```python
def solve_H(n,b):
    H = np.fromfunction(lambda i ,j : 1/ (i + j +2), (n,n))
    A = H@H - np.eye(n,n)
    x = la.solve(A,b)
    return x
```

<!-- grade_id: solve_H_type, points: 1, grade: true, solution: false -->
```python
"Check that solve_H returns the correct type."
assert type(solve_H(10,np.arange(10).reshape(10,1))) == np.ndarray, "Return type should be a NumPy array."
print("Problem 3 Test 1: Success!")
```

Output:
```
Problem 3 Test 1: Success!
```

<!-- grade_id: solve_H_shape, points: 1, grade: true, solution: false -->
```python
"Check that solve_H returns array with correct shape."
n = np.random.randint(2,10)
b = np.random.randn(n,1)
assert solve_H(n,b).shape == (n,1), "Return type should be a 2D NumPy array of size n by 1."
print("Problem 3 Test 2: Success!")
```

Output:
```
Problem 3 Test 2: Success!
```

<!-- grade_id: solve_H_values, points: 2, grade: true, solution: false -->
```python
"Check that solve_H returns the correct value."
epsilon = 10e-8
x = np.array([[-3.27992737],[-3.33209729],[-3.81807645]])
b = np.array([[0],[1],[2]])
assert la.norm(solve_H(3,b) - x) < epsilon
print("Problem 3 Test 3: Success!")
```

Output:
```
Problem 3 Test 3: Success!
```

##  Problem 4 (6 points)

Write a function called `regression_system` which takes input parameters `x`, `y` and `n` where `x` and `y` are Python lists of numbers OR 1D NumPy arrays, and `n` is a positive integer. The function returns the augemented matrix of the system of equations $X^TX \textbf{a} = X^T \textbf{y}$ representing the degree $n$ polynomial regression. The arrays `x` and `y` must be the same length $N$ where $N > n$. If `x` and `y` don't have the same length, or if $N \leq n$, print an error message and return `None`.

<!-- grade_id: regression_system, points: 0, grade: false, solution: true -->
```python
def regression_system(x,y,n):
    if (len(x) != len(y) or len(x) <= n or len(y) <= n):
        print("haha long boi")
        return None
    else:
        X = np.vander(x,N = n +1,increasing=True)
        y = np.asarray(y).reshape(len(y),1)
        Aug = np.hstack([X.T@X,X.T@y])
        return Aug
        
    
```

```python
regression_system([0,1],[2,3],1)
```

Output:
```
array([[2, 1, 5],
       [1, 1, 3]])
```

<!-- grade_id: regression_system_type, points: 1, grade: true, solution: false -->
```python
"Check that regression_system returns the correct type."
assert type(regression_system([0,1],[2,3],1)) == np.ndarray
print("Problem 4 Test 1: Success!")
```

Output:
```
Problem 4 Test 1: Success!
```

<!-- grade_id: regression_system_N, points: 1, grade: true, solution: false -->
```python
"Check that regression_system returns None when x and y are not the same length."
assert regression_system([0,1,2],[2,3],1) == None
print("Problem 4 Test 2: Success!")
```

Output:
```
haha long boi
Problem 4 Test 2: Success!
```

<!-- grade_id: regression_system_n, points: 1, grade: true, solution: false -->
```python
"Check that regression_system returns None when N >= n."
assert regression_system([0,1,2],[2,3,4],3) == None
print("Problem 4 Test 3: Success!")
```

Output:
```
haha long boi
Problem 4 Test 3: Success!
```

<!-- grade_id: regression_system_shape, points: 1, grade: true, solution: false -->
```python
"Check that regression_system returns array of the correct shape."
assert regression_system([0,1],[2,3],1).shape == (2,3)
assert regression_system([0,1,2],[2,3,4],2).shape == (3,4)
print("Problem 4 Test 4: Success!")
```

Output:
```
Problem 4 Test 4: Success!
```

<!-- grade_id: regression_system_values, points: 2, grade: true, solution: false -->
```python
"Check that regression_system returns the correct array."
epsilon = 10e-14
assert np.max(np.abs(regression_system([0,1],[2,3],1) - np.array([[2,1,5],[1,1,3]]))) < epsilon
print("Problem 4 Test 5: Success!")
```

Output:
```
Problem 4 Test 5: Success!
```

```python

```